# Oppgave 1: Hvor mange mennesker bor nært kjøpesentrene? (10 poeng)

I denne oppgaven fortsetter vi fra den forrige. Du skal gjøre en romlig kobling for å relatere data fra et befolkningsrutenett-datasett til bufferlaget opprettet i forrige øving for å finne ut hvor mange mennesker som bor i alle befolkningsrutenettceller som er innenfor 1,5 km avstand fra hvert kjøpesenter.

Bruk det samme befolkningsrutenettet som ble brukt i forelesningen og bufferlaget (shopping_centres.geojson) du lagde i forrige øving. Ha begge datafilene i samme mappe som denne notebooken, og lever både begge datafilene og notebooken til slutt.

Du kan gjerne lage flere kodeceller for å løse oppgaven, bare husk å lag de før test-cellene, og ikke slett test-cellene!

Husk å kommentere koden din underveis.

## 1a) Laste inn rutenett

Bruk det samme befolkningsrutenettet som ble brukt i forelesningen. Last det inn i en `GeoDataFrame` som har navnet `rutenett`. Reprojiser rutenettet til CRS med EPSG-kode `25832`.

<!-- grade_id: cell-7e90ec1a26581034, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

<!-- grade_id: cell-be24a61000444bc9, points: 2, grade: true, solution: false -->
```python
# TEST-CELLE
## Ikke slett denne cellen, og nye kodeceller du lager til denne deloppgaven skal komme før denne!
import geopandas
import pyproj
assert isinstance(rutenett, geopandas.GeoDataFrame), "Finner ikke GeoDataFramen 'rutenett'"
assert rutenett.crs == pyproj.CRS("EPSG:25832"), "Sjekk at GeoDataFramen har riktig CRS"

```

## 1b) Last inn bufferlag

Last inn bufferne du lagde i forrige øving inn i en `GeoDataFrame` som heter `buffer_lag`. 

Forsikre deg om at alle kjøpesentrene har en buffer, hvis ikke må du gjøre geokodingen fra Øving 4 om igjen, til du har en buffer per kjøpesenter.

Lag en `assert` som sjekker hvorvidt begge geodataframene har samme CRS.

<!-- grade_id: cell-4979b8018357a4f5, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

<!-- grade_id: cell-3967ab3d778379b5, points: 2, grade: true, solution: false -->
```python
# TEST-CELLE
## Ikke slett denne cellen, og nye kodeceller du lager til denne deloppgaven skal komme før denne!
assert isinstance(buffer_lag, geopandas.GeoDataFrame), "Finner ikke GeoDataFramen 'buffer_lag'"

```

## 1c) Utfør en romlig kobling mellom `rutenett` og `buffer_lag`

Slå sammen bufferlagets `id`-kolonne (og andre, hvis du vil) med befolkningsrutenett-dataframen, for alle befolkningsrutenettceller som er innenfor bufferområdet til hvert kjøpesenter. [Bruk en sammenkoblingstype som beholder kun rader fra begge inndataframene der den geometriske predikaten er sann](https://geopandas.org/en/stable/gallery/spatial_joins.html#Types-of-spatial-joins).

Lagre resultatet i en variabel med navnet `sammenkoblet_data`.

<!-- grade_id: cell-49d89d238209cf32, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

<!-- grade_id: cell-d9f6fec21bbe9d3c, points: 2, grade: true, solution: false -->
```python
# TEST-CELLE
## Ikke slett denne cellen, og nye kodeceller du lager til denne deloppgaven skal komme før denne!
assert isinstance(sammenkoblet_data, geopandas.GeoDataFrame), "Finner ikke GeoDataFramen 'sammenkoblet_data'"
```

## 1d) Regn ut den totale befolkningen rundt kjøpesentrene

Grupper den resulterende (sammenkoblede) dataframen etter kjøpesenternavn, og beregn summen (`sum()`) av befolkningen som bor innenfor en radius på 1,5 km rundt dem.

Skriv ut resultatene, for eksempel i formen "12345 mennesker bor innen 1,5 km fra Oslo City".

<!-- grade_id: cell-6443f2549daee070, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

<!-- grade_id: cell-d7043a1fbfe991f3, points: 2, grade: true, solution: false -->
```python
# TEST-CELLE
## Ikke slett denne cellen, og nye kodeceller du lager til denne deloppgaven skal komme før denne!

```

## 1e) Refleksjon

Godt jobbet! Du er nesten ferdig med denne ukens øvelse. Vennligst svar raskt på følgende korte spørsmål:

* Hvor utfordrende syntes du denne og forrige øvingsoppgave var (på en skala fra 1-5), og hvorfor?
* Hva var lett?
* Hva var vanskelig?

Legg til svarene dine i Markdown-cellen nedenfor:

<!-- grade_id: cell-88c2d55bfd1264a5, points: 2, grade: true, solution: true -->
SKRIV SVARENE DINE I DENNE CELLEN