# Øving 9: Nettverksanalyse (10 poeng)


Denne uken skal vi jobbe med OpenStreetMap data og utføre en nettverksanalyse i Python.

## Oppgave 1: Lese inn data

I data-mappen ligger det to filer:

- `start.csv`
- `slutt.csv`

Les de inn i to dataframes som heter henholdsvis `start` og `slutt`.

<!-- grade_id: cell-6ecbba49a2f4d187, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

<!-- grade_id: cell-3898a9dd343d8a76, points: 2, grade: true, solution: false -->
```python
# TEST-CELLE
## Ikke slett denne cellen, og nye kodeceller du lager til denne deloppgaven skal komme før denne!
```

## Oppgave 2: Lokaliser punktene

Du skal nå localisere punktene på to ulike måter:

### 2a) Visuell lokalisering

Her skal du plotte punktene på ett kart med et bakgrunnskart og se hvor punktene befinner seg. Begge datasettene er i utgangspunktet i EPSG:4326.

<!-- grade_id: cell-1a8eba087b3fc46c, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

### 2b) Geokoding

Du skal nå bruke omvendt geokoding av punktene ([se hint fra geopy-dokumentasjonen](https://geopy.readthedocs.io/en/stable/index.html?highlight=reverse#module-geopy.geocoders)) og printe ut adressene fra minst ett av lagene med punkter.

<!-- grade_id: cell-5849e2f89e2f54c5, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

## Oppgave 3: OpenStreetMap

Du skal nå hente ut data fra OpenStreetMap for området hvor punktene befinner seg. Bruk bare veinettet for bil.

Spesifiser først utstrekningen til området du skal laste inn data fra som et Shapely Polygon. Utstrekningen skal være et polygon som er lagret i variablen `utstrekning`.

- Merk at noen av rutene kan gå utover omfanget av punktdatasettene!
- Legg til litt buffer rundt punktene for å laste ned grafen, slik at du dekker alle rutene mellom punktene.
- Se hint fra denne ukens oppgavesid om hvordan du definerer utstrekningen.


<!-- grade_id: cell-e7b40a6bb6011717, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

- Bruk polygonet til laste inn det kjørbare veinettet med `.graph_from_polygon()`-funksjonen
- Husk å spesifiser riktig `network_type`
- se [osmnx-dokumentasjonen](https://osmnx.readthedocs.io/en/stable/osmnx.html#osmnx.core.graph_from_polygon) for mer detaljer
- lagre veinettet i en variabel med navn `graph`

<!-- grade_id: cell-dfdb4c23fd84e9f1, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

<!-- grade_id: cell-7bc2c163d1a34a87, points: 2, grade: true, solution: false -->
```python
# TEST-CELLE
## Ikke slett denne cellen, og nye kodeceller du lager til denne deloppgaven skal komme før denne!
```

## Oppgave 4: Reprojisering

Reprojiser dataene til en UTM-projeksjon og plot et kart med veinettet, start-punktene i en farge og slutt-punktene i en annen farge

<!-- grade_id: cell-8c0b727b217fac1e, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

<!-- grade_id: cell-c2524f893d96fade, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

## Oppgave 5: Korteste vei

Regn ut korteste vei mellom _alle_ start- og slutt-punkt ved hjelp av lengden på veisegmentene.

- For å finne de korteste rutene, bør du finne de nærmeste nodene fra grafen for både start- og destinasjonspunktene.
- Merk: hvis den nærmeste noden er den samme for både start- og destinasjonspunktet, bør du hoppe over beregningen av den korteste ruten.
- Etter ruteberegningen bør du legge til de korteste rutene som LineString-geometrier i en GeoDataFrame kalt `ruter`.
- Du bør bruke den projiserte grafen som nettverk for analysene.
- Husk å bruke informative variabelnavn og å kommentere koden din!

<!-- grade_id: cell-4d8198e95b9a5275, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

## Oppgave 6: Regn ut avstand

Regn ut avstanden til alle rutene i meter og lagre i en ny kolonne ved navn `avstand`


<!-- grade_id: cell-38b836b734d5bef7, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

<!-- grade_id: cell-747e3cce14122ece, points: 1, grade: true, solution: false -->
```python
# TEST-CELLE
## Ikke slett denne cellen, og nye kodeceller du lager til denne deloppgaven skal komme før denne!
print(f"Korteste distanse: {round(ruter['avstand'].min())} meter")
```

<!-- grade_id: cell-1dd7f7578d403ccd, points: 1, grade: true, solution: false -->
```python
# TEST-CELLE
## Ikke slett denne cellen, og nye kodeceller du lager til denne deloppgaven skal komme før denne!
print(f"Lengste distanse: {round(ruter['avstand'].max())} meter")
```

## Oppgave 7: Plot

Plot alle rutene oppå veinettet

<!-- grade_id: cell-837110a4a1ed4706, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

## Oppgave 8: Total avstand

Regn ut den totale avstanden og lagre den i en variabel ved navn `total_avstand`

<!-- grade_id: cell-977c64eb5b5d1cbb, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

<!-- grade_id: cell-b7ca6807e1344dd6, points: 0, grade: true, solution: false -->
```python
# TEST-CELLE
## Ikke slett denne cellen, og nye kodeceller du lager til denne deloppgaven skal komme før denne!
# Print løsningen din:
print(f"Den totale avstanden av alle ruter er {round((total_avstand / 1000))} kilometer")

# Sjekk om du har samme resultat som den innebygde testen.
# Ikke få panikk om du får feil her, koden din kan være riktig, men at du har brukt et litt annet nettverk i analysen tidligere i oppgaven.
assert 0.9*575 <= round(total_avstand/1000) <= 1.1*575, "Your solution doesn't match the model solution."
```

## Ferdig!

Gratulerer, du er nå ferdig med siste øvingsoppgave, da er det bare prosjektoppgaven som gjenstår!