# Oppgave 1: (10 poeng)

**1a) Les inn grunnkretser for Trondheim**

* Les inn filen `Grunnkretser 2024.geojson` og lagre den i en variabel som heter `grunnkretser`.
* Sørg for at du har lest inn en GeoDataFrame med korrekt crs.

<!-- grade_id: cell-fbdc1fa2f055ff3d, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

<!-- grade_id: cell-40945faf4743760d, points: 2, grade: true, solution: false -->
```python
# TEST-CELLE
## Ikke slett denne cellen, og nye kodeceller du lager til denne deloppgaven skal komme før denne!
import pandas
import geopandas
assert isinstance(grunnkretser, geopandas.GeoDataFrame), "Har du lest inn GeoDataFramen riktig?"
```

**1b) Les inn befolkningstall for Trondheim**

* Les inn befolkningstall for Trondheim fra filen `TRD-folkemengde.csv` i data-mappen inn i en DataFrame. NB. denne filen inneholder ikke geografisk informasjon.
* Lagre befolkningstallen i en variabel med navn `befolkning`.
* Endre navnet på kolonnen som inneholder navn på grunnkretser til å være det samme som i `grunnkretser`-GeoDataFramen. (Du kan f.eks. bruke [`.rename()`-funksjonen](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.rename.html))

<!-- grade_id: cell-430dc3a7ef9c6cad, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

<!-- grade_id: cell-a6c5b768e8a318f2, points: 2, grade: true, solution: false -->
```python
# TEST-CELLE
## Ikke slett denne cellen, og nye kodeceller du lager til denne deloppgaven skal komme før denne!
assert isinstance(befolkning, pandas.DataFrame), "Har du lest inn befolkningstallene riktig?"
```

**1c) Koble sammen grunnkretser og befolkningstall**

Du skal nå koble sammen grunnkretsene med befolkningstallene. Du skal sitte igjen med en GeoDataFrame som inneholder alle grunnkretsene og befolkningstall for alle årene. Lagre den sammenslåtte dataen i en GeoDataFrame som heter `grunnkretser_befolkning`.

<!-- grade_id: cell-4b3b55ed93eb8f08, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

<!-- grade_id: cell-2e5738086f9b59d8, points: 2, grade: true, solution: false -->
```python
# TEST-CELLE
## Ikke slett denne cellen, og nye kodeceller du lager til denne deloppgaven skal komme før denne!
assert isinstance(grunnkretser_befolkning, geopandas.GeoDataFrame), "Har du lest inn GeoDataFramen riktig?"
```

**1d) Klassifiser dataene**

Du skal nå klassifisere befolkningsdataene for 2024 i det sammenslåtte datasettet:
* Lag en `UserDefined` klassifiserer ved hjelp av [mapclassify](https://github.com/pysal/mapclassify#userdefined) og bruk klasseneskillene `250,500,1000,2000`
* Bruk klassifisereren din på kolonnen med befolkningsdata for 2024 og lagre det i en kolonnen med navn `2024_klasser`.


<!-- grade_id: cell-075711dc0f378414, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

<!-- grade_id: cell-4b93f4f0f9c08464, points: 2, grade: true, solution: false -->
```python
# TEST-CELLE
## Ikke slett denne cellen, og nye kodeceller du lager til denne deloppgaven skal komme før denne!
assert '2024_klasser' in grunnkretser_befolkning.columns, "Kolonnen '2024_klasser' finnes ikke i 'grunnkretser_befolkning'"
```

**1e) Plot de klassifiserte befolkningsdataene**

Lag et enkelt plot over befolkningsdataene basert på den nye klassifiseringen.

<!-- grade_id: cell-59ca4da5a9aaa917, points: 0, grade: false, solution: true -->
```python
# SKRIV DIN KODE HER OG FJERN LINJEN "raise NotImplementedError()"
raise NotImplementedError()
```

<!-- grade_id: cell-9ba460bbd05aec59, points: 2, grade: true, solution: false -->
```python
# TEST-CELLE
## Ikke slett denne cellen, og nye kodeceller du lager til denne deloppgaven skal komme før denne!

```

## Valgfri bonusoppgave (*2 poeng*):
Lag det samme kartet for årene 2022 og 2023 og plot et histogram over den egendefinerte klassifiseringen for hvert av de tre årene.

```python

```